```python
NAME = "Magali"
COLLABORATORS = "NA"
```

---

# CS110 Pre-class Work 11.1

The pre-class work for this session will focus on the rod cutting problem. Recall that the rod cutting problem takes as an input the length n of a rodand a table of prices $p_i$ for $i = 1,2,... n$, and one needs to determine the maximum revenue $r$ obtainable by cutting up the rod and selling the pieces. 

## Part A. 

You will implement in Python two solutions to the rod cutting problem, namely:

## Question 1. 
A recursive top-down solution to the rod cutting problem. Please complete the cut_rod function below:


<!-- grade_id: cell-d85ce9860c9ea1e5, points: 0, grade: false, solution: true -->
```python
def cut_rod(p,n):
    """
    A recursive top-down solution to the rod cutting problem as described in 
    Cormen et. al. (pp 363) This calculates the maximum revenue that could be 
    earned by cutting up a rod of length n.
    
    Inputs;:
    - p: list of floats, the prices of rods of different lengths. p[i] gives the dollars
    of revenue the company earns selling a rod of length i+1.
    - n: int, length of the rod
    
    Outputs:
    - q: float, the optimal revenue
    """
    
    if n == 0:
        return 0
    q = float('-inf')
    for i in range(n):
        q = max(q, (p[i] + cut_rod(p, n-i)))
        # print(q)
    print("final is:", q)
    return q
```

<!-- grade_id: cell-0288a4047be5088d, points: 1, grade: true, solution: false -->
```python
# price list from textbook
p = [1,5,8,9,10,17,17,20,24,30] #p[i] gives the price of length i+1

#results from textbook
r = [0,1,5,8,10,13,17,18,22,25,30]
for i in range(len(r)):
    assert(cut_rod(p,i)==r[i])
```

Output:
```
... last 1 frames repeated, from the frame below ...

<ipython-input-1-27df70571e8c> in cut_rod(p, n)
     18     q = float('-inf')
     19     for i in range(n):
---> 20         q = max(q, (p[i] + cut_rod(p, n-i)))
     21         # print(q)
     22     print("final is:", q)

RecursionError: maximum recursion depth exceeded in comparison
```

```python
# try for fewer problems / sub-problems
# price list from textbook
p = [1,5] #p[i] gives the price of length i+1

#results from textbook
r = [0,1]
for i in range(len(r)):
    print("result is:", cut_rod(p,i))
```

Output:
```
result is: 0
... last 1 frames repeated, from the frame below ...

<ipython-input-1-27df70571e8c> in cut_rod(p, n)
     18     q = float('-inf')
     19     for i in range(n):
---> 20         q = max(q, (p[i] + cut_rod(p, n-i)))
     21         # print(q)
     22     print("final is:", q)

RecursionError: maximum recursion depth exceeded in comparison
```

```python
# reset recursion limit
import sys
print(sys.getrecursionlimit())
sys.setrecursionlimit(3000)
```

Output:
```
3000
```

I'm unsure why I get this recursion depth exceeded error. I tried solving for multiple smaller subsets of problems (see above for the smallest set I tried), given that the readings mentioned that there might be a absurdly long runtime (this is in line with my analysis of the algorithm as well). 

Since this change did not work, I looked the error up, and tried to change the recursion limit, as demonstrated in the above cell.  This was not successful. 

Maybe because I mistranslated the pseudocode into Python?  I tried changing my method (my for loop) accordingly but I still get an error (assert condition is not successfully met). 

def cut_rod(p,n):
    """
    A recursive top-down solution to the rod cutting problem as described in 
    Cormen et. al. (pp 363) This calculates the maximum revenue that could be 
    earned by cutting up a rod of length n.
    
    Inputs;:
    - p: list of floats, the prices of rods of different lengths. p[i] gives the dollars
    of revenue the company earns selling a rod of length i+1.
    - n: int, length of the rod
    
    Outputs:
    - q: float, the optimal revenue
    """
    
    if n == 0:
        return 0
    q = float('-inf')
    for i in range(1,n+1):
        q = max(q, p[i] + cut_rod(p, n-i))
    return q

```python
# price list from textbook
p = [1,5,8,9,10,17,17,20,24,30] #p[i] gives the price of length i+1

#results from textbook
r = [0,1,5,8,10,13,17,18,22,25,30]
for i in range(len(r)):
    assert(cut_rod(p,i)==r[i])
```

Output:
```
... last 1 frames repeated, from the frame below ...

<ipython-input-1-27df70571e8c> in cut_rod(p, n)
     18     q = float('-inf')
     19     for i in range(n):
---> 20         q = max(q, (p[i] + cut_rod(p, n-i)))
     21         # print(q)
     22     print("final is:", q)

RecursionError: maximum recursion depth exceeded in comparison
```

```python
# see what's happening by printing instead of asserting

# price list from textbook
p = [1,5,8,9,10,17,17,20,24,30] #p[i] gives the price of length i+1

#results from textbook
r = [0,1,5,8,10,13,17,18,22,25,30]
for i in range(len(r)):
    print(cut_rod(p,i))
```

Output:
```
0
... last 1 frames repeated, from the frame below ...

<ipython-input-1-27df70571e8c> in cut_rod(p, n)
     18     q = float('-inf')
     19     for i in range(n):
---> 20         q = max(q, (p[i] + cut_rod(p, n-i)))
     21         # print(q)
     22     print("final is:", q)

RecursionError: maximum recursion depth exceeded in comparison
```

:(

## Question 2.

An optimal rod cutting solution using dynamic programming (see figure [here](https://drive.google.com/open?id=1nu9gETKX4KJCHZDi17fZLQtkDVig81Zk) for inspiration). Please complete the following two functions: 


<!-- grade_id: cell-eb49e8516ebf80aa, points: 0, grade: false, solution: true -->
```python
def extended_bottom_up_cut_rod(p,n):
    """
    Implements a bottom-up dynamic programming approach to the rod cutting problem.
    Here, "extended" means the function is geared in a way amenable to reconstructing
    an optimal solution, on top of the returned optimal value. See Cormen et al.,
    p. 369 for the implementation details.
    
    Inputs:
    - p: list of floats, the prices of rods of different lengths. p[i] gives the dollars
    of revenue the company earns selling a rod of length i+1.
    - n: int, length of the rod
    
    Outputs:
    - r: list of floats, the maximum revenues. r[i] gives the maximum revenue for a rod
    of length i. As such:
        * r[0] = 0
        * len(r) == n + 1
    - s: list of ints, the optimal sizes of the first piece to cut off. Also make sure 
    that:
        * s[0] = 0
        * len(s) == n + 1
    """
    # initialize empty arrays
    r = [0.0]*n 
    s = [0]*n
    
    # r[0] = 0
    
    for j in range(n):
        q = float('-inf')
        for i in range(j):
            if q < p[i] + r[j-i]:
                q = p[i] + r[j-i]
                s[j] = i
        r[j] = q
    return (r, s)
```

<!-- grade_id: cell-1d52c9519ccb3502, points: 0, grade: false, solution: true -->
def print_cut_rod_solution(p,n):
    """
    Gives a solution to the rod cutting problem of size n. 
    
    Inputs:
    - p: list of floats, the prices of rods of different lengths. p[i] gives the revenue (in USD, for example) the company earns selling a rod of length i+1.
    - n: int, length of the rod
    
    Outputs:
    - sol: a list of ints, indicating how to cut the rod. Cutting the rod with the lengths
    given in sol gives the optimal revenue.
        * print_cut_rod_solution(p,0) == []
    """
    (r, s) = extended_bottom_up_cut_rod(p, n)
    # n = n-1
    print("this is r:", r)
    print("this is s:", s)
    while n >= 1:
        print("n is:", n)
        print (s[n-1])
        n = n - s[n-1]

```python
def print_cut_rod_solution(p,n):
    """
    Gives a solution to the rod cutting problem of size n. 
    
    Inputs:
    - p: list of floats, the prices of rods of different lengths. p[i] gives the revenue (in USD, for example) the company earns selling a rod of length i+1.
    - n: int, length of the rod
    
    Outputs:
    - sol: a list of ints, indicating how to cut the rod. Cutting the rod with the lengths
    given in sol gives the optimal revenue.
        * print_cut_rod_solution(p,0) == []
    """
    (r, s) = extended_bottom_up_cut_rod(p, n)
    # n = n-1
    print("this is r:", r)
    print("this is s:", s)
    while n > 0:
        print (s[n])
        n = n - s[n]
```

```python
# price list from textbook
p = [1,5,8,9,10,17,17,20,24,30] #p[i] gives the price of length i+1
# Result of R and S from textbook:
R = [0,1,5,8,10,13,17,18,22,25,30]
S = [0,1,2,3,2,2,6,1,2,3,10]
# Test:
(r, s) = extended_bottom_up_cut_rod(p,10)
print(r)
print(s)
```

Output:
```
[-inf, 1.0, 6.0, 11.0, 16.0, 21.0, 26.0, 31.0, 36.0, 41.0]
[0, 0, 1, 1, 1, 1, 1, 1, 1, 1]
```

<!-- grade_id: cell-6575aa2769e0eab1, points: 1, grade: true, solution: false -->
```python
# price list from textbook
p = [1,5,8,9,10,17,17,20,24,30] #p[i] gives the price of length i+1
# Result of R and S from textbook:
R = [0,1,5,8,10,13,17,18,22,25,30]
S = [0,1,2,3,2,2,6,1,2,3,10]
# Test:
r, s = extended_bottom_up_cut_rod(p,10)
assert(r==R)
assert(s==S)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-10-0a1b9fc5af6b> in <module>
      6 # Test:
      7 r, s = extended_bottom_up_cut_rod(p,10)
----> 8 assert(r==R)
      9 assert(s==S)

AssertionError:
```

## Part B - Experiments

## Question 1.

Use the function below to generate a list of prices of length n=20 and assign that list to a new variable, `P`. You MUST use this list for parts 2 and 3 below.


```python
import numpy as np
def generate_price_list(n):
    """Generates a price list of length n
    
    Inputs:
    - n: integer, length of the list, must be positive
    
    Outputs:
    - p: list, the ordered price list for each rod cut
    """
    p = [1]
    for i in range(1,n):
        np.random.seed(0)
        p.append(np.random.randint(p[i-1]+1, i*3+1))
    return p

```

```python
price_list = generate_price_list(20)
print(price_list)
```

Output:
```
[1, 2, 3, 8, 9, 14, 15, 20, 21, 26, 27, 32, 33, 38, 39, 44, 45, 50, 51, 56]
```

## Question 2.

Time (using the time library) the implementation of `cut_rod` for different rod lengths for values of up to $n=20.$ Present your results in a plot. 


<!-- grade_id: cell-e094e913d8bc139f, points: 0, grade: true, solution: true -->
```python
import time

def time_cut_rod(p,n):
    start = time.clock() # get current processor time (in seconds)
    q = cut_rod(p, n) # call the method
    stop = time.clock() # get processor time after calling method
    time_taken = stop - start # calculate time elapsed (seconds)
    return time_taken
```

```python
t = []
n = [0, 1, 2, 3, 4 ,5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20]

# call the method
for i in range(21):
    print("time taken for n =", i, "is:")
    time_taken = time_cut_rod(price_list, i)
    print(time_taken)
    t.append(time_taken)
```

```python
# plot
import matplotlib.pyplot as plt

time_plot = plt.plot(n, t)
a = plt.ylabel("time taken to implement cut_rod() in seconds")
b = plt.xlabel("length of the rod")
```

## Question 3.

Time (using the time library) the implementation of `extended_bottom_up_cut_rod` for different rod lengths for values of up to $n=20$. Add the curve corresponding to this algorithm to the previous plot.

```python
def time_extended_bottom_up_cut_rod(p,n):
    start = time.clock() # get current processor time (in seconds)
    q = cut_rod(p, n) # call the method
    stop = time.clock() # get processor time after calling method
    time_taken = stop - start # calculate time elapsed (seconds)
    return time_taken
```

<!-- grade_id: cell-210aa3d5a2743bd8, points: 0, grade: true, solution: true -->
```python
t = []
n = [0, 1, 2, 3, 4 ,5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20]

# call the method
for i in range(21):
    print("time taken for n =", i, "is:")
    time_taken = time_extended_bottom_up_cut_rod(price_list, i)
    print(time_taken)
    t.append(time_taken)
```

```python
ex_line = plt.plot(t, color = "red")
time_plot + ex_line + a + b 
plt.show()
```

## Question 4.

Contrast both curves. What can you conclude?

<!-- grade_id: cell-0e62cc5b19840488, points: 0, grade: true, solution: true -->
Unfortunately, I am unable to do so here because my original functions cannot be computed because of the maximum recursion depth error.  However, I imagine that the extended bottom-up approach is more efficient than the simple recursion one as the answers of the sub-problems that have already been computed are re-used instead of re-calculated.

## Part C - After completing your experiments above, answer the following questions:

## Question 1.
Estimate how big $n$ has to be before the non-dynamic-programming version will take a million years to finish. Call this value N.

<!-- grade_id: cell-a11f4fd6a03e2cf5, points: 0, grade: true, solution: true -->
I am unable to know how long it takes for the non-dynamic-programming version to finish for a small n above 0 (e.g. 5) as explained above. However this would be my approach:

1. Find a million years in seconds: 3.1536 x 10^13. Round off to 3 x 10^13. 

2. Find out how time scales with n. I am assuming this will be exponential.  With lack of data, let's assume make a modest assumption of 2^n (of course, 2 should be replaced with an accurate estimation based on what the plots show). 

3. Plugging these estimations in: 

    3 x (10^13) = 2^n 

    (We could ignore the constant 3 in this exponential growth if we wanted.)

    n = log2(3 x 10^13).
    n = 44.

    I estimate that N should be equal to 44, under my time assumption in 2 (which should be modified according to what the graph displays).  

## Question 2.
Estimate (or time) how long it takes to evaluate `extended_bottom_up_cut_rod(p,N)`, where `N` is the value you got from the previous question. 

<!-- grade_id: cell-89ed2dd6f3310c96, points: 0, grade: true, solution: true -->
```python
N = 44 # define the value from above
n_time = time_extended_bottom_up_cut_rod(p,N)
```

## [Optional]  Question 3. 
Do you notice anything peculiar about the solution for large n?

<!-- grade_id: cell-fd4c9c91a42f178a, points: 0, grade: true, solution: true -->
Is it flattening out?

# Personal exploration of growth

```python
def plot_scaling(n, m):
    result_power = []
    result_exp = []
    for i in n:
        # print("i:", i)
        # print(i**10)
        result_power.append(i**m)
        # print(10**i)
        result_exp.append((m**i))

    plt.plot(n, result_power, color = 'red')
    plt.plot(n, result_exp, color = 'blue')
```

```python
plot_scaling([0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15], 10)
```

```python
plot_scaling([0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 5)
plot_scaling(n, 8)
```

```python
plot_scaling([0, 1, 2, 3, 4, 5], 5)

```