```python
NAME = "Magali"
COLLABORATORS = "NA"
```

---

# CS110 Pre-class Work 11.2

## Part A.  Currency trading (Slightly simplified version of 15.3-6 from Cormen et al.)

Imagine that you wish to exchange one currency for another. You realize that instead of directly exchanging one currency for another, you might be better off making a series of trades through other currencies, winding up with the currency you want.

Suppose that you can trade n different currencies, numbered $1,2,… ,n$, where you start with currency 1 and wish to wind up with currency $n$. You are given, for each pair of currencies $i$ and $j$ , an exchange rate $r_{ij}$ , meaning that if you start with $d$ units of currency $i$ , you can trade for $dr_{ij}$ units of currency $j$. Note that the total number of trades allowed is limited to $n$.

Assuming there is no commission, propose an algorithm (described in prose) to solve this problem using either a memoization or bottom-up strategy. **Note that we will be working on the Python implementation in class.**

Click [here](https://drive.google.com/open?id=1L8Fjo1Xt8sltab-tz3m91eTuiE9LYWF7) for some example data.

<!-- grade_id: cell-3f8a8808c720baf2, points: 0, grade: true, solution: true -->
This is an optimization problem where we seek to find the optimal way trades.  Hence we need to convert each currency into another currency and find the maximum value of each of these subproblems, and then recombine them accordingly.  This can be achieved through a bottom-up dynamic programming algorithm. We solve all of the smallest subproblem conversion first. We convert from currency 1 to 2 or 3 or 4.... all the way up to n and repeat this each time the currency has been exchanged, except we do not repeat past cases. We can then choose each optimal subsolution and compose the optimal solution from this. 

Another way to look at this is that theoretically the exchange from currency X to Y and from currency Y and X should add up to 2. So if a currency adds up to less than that, then it is an indication that it is either overvalued or undervalued. This can inform our choice of currencies for conversion. More thought on how this could be implemented algorithmically is needed.

## Part B - Money game (Solution)
Consider a row of n coins of values $v_1, v_2,...,v_n$, **where $n$ is even**. We play a game against an opponent by alternating turns. In each turn, a player selects either the first or last coin from the row, removes it from the row permanently, and receives the value of the coin. Determine the maximum possible amount of money we can definitely win if we move first.

For example, consider the game:

\$2, \$10, \$1, \$5

By moving first and playing optimally one can be guaranteed of \$15. The first move is to take \$5. This forces your opponent to take either \$2 or \$1, and then allows you to take \$10. Assume that the opposing player also plays optimally (i.e., minimizing the gain of the first-move player.)

As a hint, we will provide the recurrence for the solution to this problem. Suppose $c(i,j)$ is the maximum possible amount of money the first player can win for the row of coins of value $v_i, v){i+1},...,v_j$ (so to solve our problem, we need to compute $c(1,n)$. Then:

$$c(i,j)=max[v_i+min(c(i+2,j), c(i+1,j-1)), v_j+min(c(i,j-2), c(i+1,j-1))]$$

You are also encouraged to watch [this video](https://www.youtube.com/watch?v=KnP8_L13xW4&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5) if you need some more hints.

## Question 1. 
Complete the following function to solve the game using a bottom-up strategy, assuming that the opposing player also plays optimally. 


```python
def cut_array(A):
    print(A)
    last_i = len(A)-1
    # print("last index:", last_i)
    if A[0] >= A[last_i]:
        p = A[0]
        A = A[1:]
    else:
        p = A[last_i]
        A =  A[:last_i]
    return A, p
```

print(cut_A([5, 1, 2, 3]))

```python
def bottom_up_coin_game(A):
    """
    Returns the maximum possible amount of money the first-move player can win,
    given an array of coin values. 
    
    Inputs:
    - A: list of floats, values of the coins, of even length. 
    
    Outputs:
    - max_val: float, maximum possible amount of money that can be won by the first 
    player. max_val is None when the length of list A is odd.
    """
    
    if len(A) % 2 != 0:
        return None
    if len(A) == 0:
        return 0
    
    ttl_made = []
    t = 0
    
    for j in range(0, len(A)):
        A, coin = cut_array(A)
        t = t+1
        if t % 2 != 0: # assuming I start
            ttl_made.append(coin)
    
    max_value = sum(ttl_made)
    print("what I got as coins, as a first player:", ttl_made)
    print("total made:", max_value)
    
    return max_value
```

<!-- grade_id: cell-f7a5d6c26e95dab2, points: 1, grade: true, solution: false -->
```python
import numpy as np
assert(bottom_up_coin_game([2, 10, 1, 5]) == 15)  
assert(bottom_up_coin_game([2, 10, 1, 5, 5]) is None)  
assert(bottom_up_coin_game([]) == 0)  
```

Output:
```
[2, 10, 1, 5]
last index: 3
[2, 10, 1]
last index: 2
[10, 1]
last index: 1
[1]
last index: 0
what I got, as a first player: [5, 10]
```

<!-- grade_id: cell-4014ae901b66304b, points: 0, grade: false, solution: true -->
```python
def bottom_up_coin_game(A):
    """
    Returns the maximum possible amount of money the first-move player can win,
    given an array of coin values. The function solves this using a bottom-up 
    approach.
    
    Inputs:
    - A: list of floats, values of the coins, of even length. 
    
    Outputs:
    - max_val: float, maximum possible amount of money that can be won by the first 
    player. max_val is None when the length of list A is odd.
    """
    
    # store = [][]
    w, h = len(A), len(A)
    store = [[0 for x in range(w)] for y in range(h)] 
    
    n = len(A)
    for i in range(0, n):
        a = n
        sol = float('-inf')
        temp_A = A
        for j in range(0, a):
            print("\nturn:", j)
            temp_A, sol = cut_array(temp_A)
            store[j][j+i] = sol  # modify this more to be along the lines of max(sol, ...)
            print("the pick is:", sol)
            print("the array is now:", A)
            a = a-1
        n = n-1
    
    print("final picks:", store)
    max_val = max(store)
    
    return max_val
```

```python
print(coin_game([2, 10, 1, 5]))
```

Output:
```
[2, 10, 1, 5]
last index: 3
[2, 10, 1]
last index: 2
[10, 1]
last index: 1
[1]
last index: 0
what I got: [5, 10]
15
```

```python
print("\nthe maximum is:", bottom_up_coin_game([2, 10, 1, 5]))
```

Output:
```
turn: 0
[2, 10, 1, 5]
last index: 3
the pick is: 5
the array is now: [2, 10, 1, 5]

turn: 1
[2, 10, 1]
last index: 2
the pick is: 2
the array is now: [2, 10, 1, 5]

turn: 2
[10, 1]
last index: 1
the pick is: 10
the array is now: [2, 10, 1, 5]

turn: 3
[1]
last index: 0
the pick is: 1
the array is now: [2, 10, 1, 5]

turn: 0
[2, 10, 1, 5]
last index: 3
the pick is: 5
the array is now: [2, 10, 1, 5]

turn: 1
[2, 10, 1]
last index: 2
the pick is: 2
the array is now: [2, 10, 1, 5]

turn: 2
[10, 1]
last index: 1
the pick is: 10
the array is now: [2, 10, 1, 5]

turn: 0
[2, 10, 1, 5]
last index: 3
the pick is: 5
the array is now: [2, 10, 1, 5]

turn: 1
[2, 10, 1]
last index: 2
the pick is: 2
the array is now: [2, 10, 1, 5]

turn: 0
[2, 10, 1, 5]
last index: 3
the pick is: 5
the array is now: [2, 10, 1, 5]
final picks: [[5, 5, 5, 5, 0, 0, 0, 0], [0, 2, 2, 2, 0, 0, 0, 0], [0, 0, 10, 10, 0, 0, 0, 0], [0, 0, 0, 1, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0, 0, 0]]

the maximum is: [5, 5, 5, 5, 0, 0, 0, 0]
```

n = 4
for i in range(0, n):
    a = n
    for j in range(0, a):
        print(j, j+i)
        a = a-1
    n = n-1

```python
import numpy as np
assert(bottom_up_coin_game([2, 10, 1, 5]) == 15)  
assert(bottom_up_coin_game([2, 10, 1, 5, 5]) is None)  
assert(bottom_up_coin_game([]) == 0)  
```

## [Optional]  Question 2. 

Complete the function `print_strategy` to print out a solution to the coin game. Completing the helper function `extended_bottom_up_coin_game` to use in `print_strategy` is optional.

<!-- grade_id: cell-4583755bc245e1e9, points: 0, grade: false, solution: true -->
```python
def print_strategy(A):
    """
    Print coins to pick for the first player
    
    Inputs:
    - A: list of floats, values of the coins, of even length.
    
    Outputs:
    - max_val: float, maximum possible amount of money that can be won by the first 
    player. 
        * max_val is None if the length of list A is odd.
        * max_val is 0 if list A is empty (no coins)
    - out: list of values of the coins that the first player picks in that
    order.
        * out is None if max_val is 0 or None.
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-75ce551c5185e374, points: 0, grade: true, solution: false -->
```python
assert(print_strategy([4,6,5,20,4,6,4,8])[0] == 40)
assert(print_strategy([4,6,5,20,4,6,4,8])[1] == [8, 6, 6, 20])
```