```python
NAME = "Magali"
COLLABORATORS = "NA"
```

---

# CS110 Pre-class Work 13.1

## Question 1.

Complete the following functions, following the algorithms in Cormen et al.

<!-- grade_id: cell-542b2dce56293aef, points: 0, grade: false, solution: true -->
```python
def lcs_length(x, y):
    """
    Computes the length of an LCS of strings x and y.
    
    Inputs:
    - x, y: strings
    
    Outputs:
    - c: a list of lists of ints OR a numpy array. c[i,j] contains the 
    length of a LCS of x[:i] and y[:j]
    - b: a list of lists of strings OR a numpy array, containing the information
    used for LCS reconstruction (See Cormen et al.) Use "N" (North), "NW" 
    (North West), and "W" (West) that correspond to the directions of the arrows 
    used in Cormen et al.
    """
    # put into lists
    x = list(x)
    y = list(y)
    
    # calculate length of strings
    m = len(x)
    n = len(y)
    
    # create the tables, initiate with 0s
    b = [[0 for ni in range(n)] for mi in range(m)]
    c = [[0 for ni in range(n+1)] for mi in range(m+1)]
    
    # loop through the strings and populate the table
    for i in range(1, m):
        for j in range(1, n):
            if x[i] == y[j]:
                c[i][j] = c[i-1][j-1] + 1
                b[i][j] = "NW"
            elif c[i-1][j] >= c[i][j-1]:
                c[i][j] = c[i-1][j]
                b[i][j] = "N"
            else:
                c[i][j] = c[i][j-1]
                b[i][j] = "W"
                
    return c, b
```

<!-- grade_id: cell-6b825bf900c2cf0b, points: 0, grade: false, solution: true -->
```python
def print_lcs(b,x,i,j,lcs):
    """
    Finds a LCS.
    
    Inputs:
    - b: a list of lists of strings OR a numpy array, returned by lcs_length
    - x: string, an input to lcs_length
    - i, j: ints. print_lcs(b,x,i,j) returns a lcs of x[:i] and y[:j], where y
    is an input to lcs_length.
    
    Outputs:
    - lcs: list of strings, representing a LCS of x and y
    - length: int, the length of the LCS
    
    You can choose to actually PRINT OUT the LCS or not using the print function.
    
    """
    
    if (i==0 and j==0):
        return 0 
    if b[i][j] == "NW":
        print_lcs(b, x, i-1, j, lcs)
        lcs = lcs.append(x[i])
        print(lcs)
    elif b[i][j] == "N":
        print_lcs(b, x, i-1, j, lcs)
    else:
        print_lcs(b, x, i, j-1, lcs)
        
    return lcs, len(lcs)
```

__Testing__

```python
x, y = 'ambgdec', 'aubyci'
c, b = lcs_length(x, y)
print("\n c is:", c)
print("\n b is:", b)
```

Output:
```
c is: [[0, 0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0, 0], [0, 0, 1, 1, 1, 1, 0], [0, 0, 1, 1, 1, 1, 0], [0, 0, 1, 1, 1, 1, 0], [0, 0, 1, 1, 1, 1, 0], [0, 0, 1, 1, 2, 2, 0], [0, 0, 0, 0, 0, 0, 0]]

 b is: [[0, 0, 0, 0, 0, 0], [0, 'N', 'N', 'N', 'N', 'N'], [0, 'N', 'NW', 'W', 'W', 'W'], [0, 'N', 'N', 'N', 'N', 'N'], [0, 'N', 'N', 'N', 'N', 'N'], [0, 'N', 'N', 'N', 'N', 'N'], [0, 'N', 'N', 'N', 'NW', 'W']]
```

```python
a, b = print_lcs(b, x, len(x)-1, len(y)-1, [''])
```

Output:
```
None
<ipython-input-17-b59f3a57adee> in print_lcs(b, x, i, j, lcs)
     26         print_lcs(b, x, i-1, j, lcs)
     27     else:
---> 28         print_lcs(b, x, i, j-1, lcs)
     29 
     30     return lcs, len(lcs)

<ipython-input-17-b59f3a57adee> in print_lcs(b, x, i, j, lcs)
     28         print_lcs(b, x, i, j-1, lcs)
     29 
---> 30     return lcs, len(lcs)

TypeError: object of type 'NoneType' has no len()
```

<!-- grade_id: cell-e06e8f09bb8a5063, points: 1, grade: true, solution: false -->
```python
import numpy as np
x, y = 'ambgdec', 'aubyci'
c, b = lcs_length(x, y)
assert(print_lcs(b,x,len(x)-1,len(y)-1)[0] == ['a', 'b', 'c'])
assert(print_lcs(b,x,len(x)-1,len(y)-1)[1] == 3)

x, y = 'xyqwsssazdesaqqf', 'xoppoypllzookjdef'
c, b = lcs_length(x, y)
assert(print_lcs(b,x,len(x)-1,len(y)-1)[0]  == ['x', 'y', 'z', 'd', 'e', 'f'])
assert(print_lcs(b,x,len(x)-1,len(y)-1)[1]  == 6)

```

Output:
```
TypeError                                 Traceback (most recent call last)
<ipython-input-20-7db5275a991e> in <module>
      2 x, y = 'ambgdec', 'aubyci'
      3 c, b = lcs_length(x, y)
----> 4 assert(print_lcs(b,x,len(x)-1,len(y)-1)[0] == ['a', 'b', 'c'])
      5 assert(print_lcs(b,x,len(x)-1,len(y)-1)[1] == 3)
      6 

TypeError: print_lcs() missing 1 required positional argument: 'lcs'
```

## Question 2. (Adapted from Exercise 15-4.1 Cormen et al.)
Use the functions built in Question 1 to find the LCS of `'10010101'` and `'010110110'`. You should store the list that represents the LCS you found in a variable named `lcs_q2`

<!-- grade_id: cell-5d742a82dde1229f, points: 0, grade: true, solution: true -->
```python
lcs_q2, len_lcs_q2 = print_lcs('10010101', '010110110')
print(lcs_q2)
```

```python
lcs_q2 == ['1', '0', '1']
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-7-7687be3231f5> in <module>
----> 1 lcs_q2 == ['1', '0', '1']

NameError: name 'lcs_q2' is not defined
```

## Question 3. (Adapted from Exercise 15-4.5 Cormen et al.)
Complete the following function, making use of `lcs_length` and `print_lcs`. 



<!-- grade_id: cell-429c944f34a54c91, points: 0, grade: true, solution: true -->
```python
def lmis(lst):
    """
    Finds the Longest Monotonically Increasing Subsequence (LMIS) of a list 
    (lst) of n numbers in O(n^2) time. Note that a monotonically increasing 
    sequence is a sequence of numbers such that: a_1 <= a_2 <= ... <= a_n .
    
    Inputs:
    - lst: a list of ints
    
    Outputs:
    - out_lst: a list of ints, a longest monotonically increasing subsequence
    of lst
    """
    
    # get only monotically increasing sequences
    milst = sorted(lst)
    print(milst)
    
    # call the method
    out_lst = print_lcs(lst, milst, len(lst), len(milst), [])
    
    return out_lst
```

<!-- grade_id: cell-542b2dce56293aef, points: 0, grade: false, solution: true -->
```python
def lcs_length(x, y):
    """
    Computes the length of an LCS of strings x and y.
    
    Inputs:
    - x, y: strings
    
    Outputs:
    - c: a list of lists of ints OR a numpy array. c[i,j] contains the 
    length of a LCS of x[:i] and y[:j]
    - b: a list of lists of strings OR a numpy array, containing the information
    used for LCS reconstruction (See Cormen et al.) Use "N" (North), "NW" 
    (North West), and "W" (West) that correspond to the directions of the arrows 
    used in Cormen et al.
    """
    # put into lists
    x = list(x)
    y = list(y)
    
    # calculate length of strings
    m = len(x)
    n = len(y)
    
    # create the tables, initiate with 0s
    b = [[0 for ni in range(n)] for mi in range(m)]
    c = [[0 for ni in range(n+1)] for mi in range(m+1)]
    
    # loop through the strings and populate the table
    for i in range(1, m):
        for j in range(1, n):
            if x[i] == y[j]:
                c[i][j] = c[i-1][j-1] + 1
                b[i][j] = "NW"
            elif c[i-1][j] >= c[i][j-1]:
                c[i][j] = c[i-1][j]
                b[i][j] = "N"
            else:
                c[i][j] = c[i][j-1]
                b[i][j] = "W"
                
    return c, b
```

<!-- grade_id: cell-6b825bf900c2cf0b, points: 0, grade: false, solution: true -->
```python
def print_lmcs(b,x,i,j,lcs):
    """
    Finds a LMCS.
    
    Inputs:
    - b: a list of lists of strings OR a numpy array, returned by lcs_length
    - x: string, an input to lcs_length
    - i, j: ints. print_lcs(b,x,i,j) returns a lcs of x[:i] and y[:j], where y
    is an input to lcs_length.
    
    Outputs:
    - lcs: list of strings, representing a LCS of x and y
    - length: int, the length of the LCS
    
    You can choose to actually PRINT OUT the LCS or not using the print function.
    
    """
    
    if (i==0 and j==0):
        return 0 
    if b[i][j] == "NW":
        print_lcs(b, x, i-1, j, lcs)
        lcs = lcs.append(x[i])
        print(lcs)
    elif b[i][j] == "N":
        print_lcs(b, x, i-1, j, lcs)
    else:
        print_lcs(b, x, i, j-1, lcs)
        
    return lcs, len(lcs)
```

```python
lmis([5, 1, 0, 4, 2, 6, 7, 9])
```

Output:
```
[0, 1, 2, 4, 5, 6, 7, 9]
<ipython-input-42-bc1b075aa3e1> in lmis(lst)
     18 
     19     # call the method
---> 20     out_lst = print_lcs(lst, milst, len(lst), len(milst), [])
     21 
     22     return out_lst

<ipython-input-17-b59f3a57adee> in print_lcs(b, x, i, j, lcs)
     19     if (i==0 and j==0):
     20         return 0
---> 21     if b[i][j] == "NW":
     22         print_lcs(b, x, i-1, j, lcs)
     23         lcs = lcs.append(x[i])

IndexError: list index out of range
```

<!-- grade_id: cell-f8b218898fb43d19, points: 0, grade: true, solution: false -->
```python
assert(lmis([5, 1, 0, 4, 2, 6, 7, 9]) == [1, 4, 6, 7, 9])
assert(lmis([6, -1,4,5,5,7,-6,2]) == [-1, 4, 5, 5, 7])
```

Output:
```
None
<ipython-input-40-3dc359f8c8a9> in <module>
----> 1 assert(lmis([5, 1, 0, 4, 2, 6, 7, 9]) == [1, 4, 6, 7, 9])
      2 assert(lmis([6, -1,4,5,5,7,-6,2]) == [-1, 4, 5, 5, 7])

<ipython-input-37-bf72260dba7b> in lmis(lst)
     18 
     19     # call the method
---> 20     out_lst = print_lcs(lst, milst, len(lst), len(milst), [])
     21 
     22     return out_lst

TypeError: object of type 'NoneType' has no len()
```

```python

```

```python

```