Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Magali"
COLLABORATORS = "NA"
```

---

This is a Python implementation of the algorithm that finds the maximum subarray. I also compare the running times of various algorithms we have covered so far, using big-Θ and big-O notations.

# CS110 Pre-class Work 2.2

## Question 1 (Exercise 3.1-3 of Cormen, et al. )
Explain why the statement, "The running time of algorithm A is at least $O(n^2)$," is meaningless.


<!-- grade_id: cell-437d61c1420d4f99, points: 0, grade: true, solution: true -->
The Big-O notation describes the worst-case scenario (the upper bound). Consequently, the statement should be: "The running time of algorithm A is *at most* O(n^2)." Otherwise, it is meaningless because the symbolism represented by Big-O ("at most") and the wording of the sentence ("at least") directly contradict each other.

Meanwhile, Big-Omega describes the best-case scenario.  Consequently, the above statement could be: "The running time of algorithm A is at least Big-Omega(n^2)."

## Question 2 (Exercise 3.1-4 of Cormen, et al. )

Is $2^{n+1}=O(2^n)$? Is $2^{2n}=O(2^n)$?

<!-- grade_id: cell-6a0c4ee19dadfddf, points: 0, grade: true, solution: true -->
Yes, 2^(n+1) = O(2^n).  2^(n+1) can be rewritten as 2^n * 2^1 or 2^n * 2.  'O' encompases the constants, hence the prior break down can be rewritten as O(2^n). 


No, 2^(2n) != O(2^n). The exponent 2 increases the outcome asymptotically and consequently, without it, 2^n represents a different value/increase. (Though I understand this, I'm curious to hear other explanations of why this is so...)

## Question 3.
Write a function in Python that solves the maximum-subarray problem using a brute-force approach. Your Python function must:
* Take as Input an array/list  of numbers
* Produce the following Output: 
    * the start and end indices of the subarray containing the maximum sum.
    * value of the maximum subarray (float)


<!-- grade_id: cell-527e6532d6992fd0, points: 0, grade: false, solution: true -->
```python
def bruteforce_max_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    
    # Check if list is NOT empty, made of floats
    if len(A) == 0:
        return ("Please provide a NON-EMPTY list.")
    if type(A[1]) != float:
        return("Please provide a list of FLOATS (not integers).")
    
    # Initiate values
    start = 0
    end = len(A)-1
    max_s = A[start]  # Sets the highest sum as the first value on its own 
    
    # Go through array to find optimal start
    for i in range(0, len(A)-1): # Loops through the array, using each element as the start one-by-one
        for j in range(i+1, len(A)-1): # Loops through the array, using each element after the start element 
                                        # as the end, one-by-one
            if A[i] >= max_s: # Compares the sum of current subarray with that of the current maximum subarray
                max_s = A[i] # If it's bigger, sets the maximum sum as that from the current subarray
                start = i # Specifies the new start index for the new-found maximum subarray
                end = j-1 # Specifies the end index for the new-found maximum subarray 
                            # (j-1 because A[i] is still the sum without the current value at the j index added)
            A[i] = A[i] + A[j]  # Adds the next index value to the current sum
                                # Putting the summation at the end allows for comparisons of one-length max. subarrays 
    
    # Output tuple: start index, end index, value of subarray
    return (start, end, max_s)
    
    raise NotImplementedError()
```

<!-- grade_id: cell-a28a56466c9537ff, points: 1, grade: true, solution: false -->
```python
assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-30-5ddab6f2f574> in <module>
----> 1 assert(bruteforce_max_subarray([-2,1,-1,2,-5]) == (1, 3, 2))
      2 assert(bruteforce_max_subarray([-2, -5, 6, -2, -3, 1, 5, -6]) == (2, 6, 7))

AssertionError:
```

## Question 4. 
Test your Python maximum-subarray function using the following input list (from Figure 4.3 of Cormen et al.):  
`A = [13, -3, -25, 20, -3, -16, -23, 18, 20, -7, 12, -5, -22, 15, -4, 7] `

If your Python implementation is correct, your code must return: 
* 43 - which is the answer to the maximum subarray problem, and 
* <7, 10> -the start and the end indices of the max subarray. 



<!-- grade_id: cell-e4a632ce0edc1697, points: 0, grade: true, solution: true -->
```python
A = [13.0, -3.0, -25.0, 20.0, -3.0, -16.0, -23.0, 18.0, 20., -7.0, 12.0, -5.0, -22.0, 15.0, -4.0, 7.0]
bruteforce_max_subarray(A)

# raise NotImplementedError()
```

Output:
```
(7, 10, 43.0)
```

## Question 5. Asymptotic notation. 
Complete the following table using the asymptotic notation that best describes the problem. For example, if both $O(n^3)$ and $O(n)$ are possible for an algorithm, the answer is $O(n)$ because the function $f(n) = O(n)$ provides a tighter and more accurate fit; if both $O(n)$ and $\Theta(n)$ are possible, the correct answer is $\Theta(n)$ because $\Theta(n)$ provides both information about the upper and lower bound, thus it is more accurate than $O(n)$.

You should copy the following table and paste and edit it in the cell below. 

Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort |  |
Selection sort |  |
Bubble sort |  | 
Finding maximum subarray |  |

<!-- grade_id: cell-247c51df276b622e, points: 0, grade: true, solution: true -->

Algorithm | Big Oh ($O$) | Big Theta ($\Theta$)
--- | --- | ---
Insertion sort |     | $\Theta$(n^2)
Selection sort |     | $\Theta$(n^2)
Bubble sort |     | $\Theta$(n^2)
Finding maximum subarray |     | $\Theta$(n^2)

According to the description above we will always favor knowing Big-Theta because it describes both the upper and lower bound.  I've specified Big-Oh as well, for future reference.

Assuming that we are finding the maximum subarray using brute-force.

Source: https://www.geeksforgeeks.org/time-complexities-of-all-sorting-algorithms/

## [Optional] Question 6. 
How can you change this code to make it find the minimum-subarray?

<!-- grade_id: cell-f4553f71858d1bc4, points: 0, grade: false, solution: true -->
```python
def bruteforce_min_subarray(A):
    """Implements brute-force minimum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the min subarray
    - the end index of the min subarray
    - the value of the min subarray
    """
   
    # Check if list is NOT empty, made of floats
    if len(A) == 0:
        return ("Please provide a NON-EMPTY list.")
    if type(A[1]) != float:
        return("Please provide a list of FLOATS (not integers).")
    
    # Initiate values
    start = 0
    end = len(A)-1
    min_s = A[start]  # Sets the highest sum as the first value on its own 
    
    # Go through array to find optimal start
    for i in range(0, len(A)-1): # Loops through the array, using each element as the start one-by-one
        for j in range(i+1, len(A)-1): # Loops through the array, using each element after the start element 
                                        # as the end, one-by-one
            if A[i] <= min_s: # Compares the sum of current subarray with that of the current maximum subarray
                min_s = A[i] # If it's bigger, sets the maximum sum as that from the current subarray
                start = i # Specifies the new start index for the new-found maximum subarray
                end = j-1 # Specifies the end index for the new-found maximum subarray 
                            # (j-1 because A[i] is still the sum without the current value at the j index added)
            A[i] = A[i] + A[j]  # Adds the next index value to the current sum
                                # Putting the summation at the end allows for comparisons of one-length max. subarrays 
    
    # Output tuple: start index, end index, value of subarray
    return (start, end, min_s)

    raise NotImplementedError()
```

<!-- grade_id: cell-4263c7494a5f09d3, points: 0, grade: true, solution: false -->
```python
assert(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
assert(bruteforce_min_subarray([1]*10)[2] == 1)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-33-c46b6dc5626e> in <module>
----> 1 assert(bruteforce_min_subarray([1]*10)[0] == bruteforce_min_subarray([1]*10)[1])
      2 assert(bruteforce_min_subarray([1]*10)[2] == 1)

AssertionError:
```

#### Testing

```python
A = [13.0, -3.0, -25.0, 20.0, -3.0, -16.0, -23.0, 18.0, 20., -7.0, 12.0, -5.0, -22.0, 15.0, -4.0, 7.0]
bruteforce_min_subarray(A)
```

Output:
```
(1, 6, -50.0)
```

```python

```