Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Magali"
COLLABORATORS = "NA"
```

---

# CS110 Pre-class Work 3.1

## Question 1. 
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def bruteforce_max_subarray(A):
    """Implements brute-force maximum subarray finding.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: A tuple of
    - the start index of the max subarray
    - the end index of the max subarray
    - the value of the maximum subarray
    """
    
    # Check if list is NOT empty, made of floats, as per instructions
    if len(A) == 0:
        return ("Please provide a NON-EMPTY list.")
    if type(A[1]) != float:
        return("Please provide a list of FLOATS (not integers).")
    
    # Initiate values
    start = 0
    end = len(A)-1
    max_s = A[start]  # Sets the highest sum as the first value on its own 
    
    # Go through array to find optimal start
    for i in range(0, len(A)-1): # Loops through the array, using each element as the start one-by-one
        for j in range(i+1, len(A)-1): # Loops through the array, using each element after the start element 
                                        # as the end, one-by-one
            if A[i] >= max_s: # Compares the sum of current subarray with that of the current maximum subarray
                max_s = A[i] # If it's bigger, sets the maximum sum as that from the current subarray
                start = i # Specifies the new start index for the new-found maximum subarray
                end = j-1 # Specifies the end index for the new-found maximum subarray 
                            # (j-1 because A[i] is still the sum without the current value at the j index added)
            A[i] = A[i] + A[j]  # Adds the next index value to the current sum
                                # Putting the summation at the end allows for comparisons of one-length max. subarrays 
    
    # Output tuple: start index, end index, value of subarray
    return (start, end, max_s)
    
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-5-0e754dec4184> in <module>
     37     return (start, end, max_s)
     38 
---> 39 raise NotImplementedError()

NotImplementedError:
```

```python
# Test method
A = [-2.0, -3.0, 4.0, -1.0, -2.0, 1.0]
bruteforce_max_subarray(A)
```

Output:
```
(2, 2, 4.0)
```

## Question 2. 
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
Yes, because the algorithm currently calculates the sum size for each index one by one.  If the appended number is large enough, it might make sense to include it (and its left-hand neighbors preceding the original end) in the subarray.

## Question 3. 

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
False.  If it is different it will include B because this is the only thing that differs between array A and B and which could bring about this change. 

## Question 4.
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Inputs:
    - x: a NON-EMPTY list of numbers (e.g., the array B in the first two questions above). 
        * If x has 1 element: returns the value of the element regardless of the value of mx
    - mx: the maximum subarray of x excluding its last element (i.e., compute the 
    maximum subarray of the input array x considering only its first len(x) - 1 elements)
    
    Output: The maximum subarray of the array x.
    
    For example, using the array B in question 2, the result of incremental_max_subarray(B, 4) 
    is 10 (10 = 8 + 1 - 2 -1 + 4). 
    """
    
    
    raise NotImplementedError()
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

<!-- grade_id: cell-bacebe71191cbd0f, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A to produce the value of the maximum
    subarray of A.
    
    Inputs:
    - A: a NON-EMPTY list of floats
    
    Outputs: flaot, the sum of the maximum subarray of A
    """
    biggest = 0
    for i in range(len(A), 0, -1):
        new = incremental_max_subarray(A)
        if new > biggest:
            biggest = new
    return biggest
    raise NotImplementedError()

```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2.
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
It will depend on the size of the array.  However, in general, it is not, because it is a brute-force approach (O(N^2), unlike the divide-and-conquer approach whose recurrence relation is lesser.

```python

```