Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Magali"
COLLABORATORS = "NA"
```

---

# CS110 Pre-class Work 3.2

## Question 1.
Given the array `H=[39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]`, perform the following operations:
1. Draw the corresponding binary tree of H. Is the binary tree a valid max heap? Explain your answer.
2. Using as a model the drawing examples illustrated in Figure 6.2 of Cormen et al.,  draw a step-by-step transformation of the array above into a valid max heap. 
3. Now that you have obtained a valid max heap, write out the corresponding array that stores the valid max-heap.

Use as many cells as you wish for this question.

<!-- grade_id: cell-b1f821b034b44619, points: 0, grade: true, solution: true -->
1. No it is not. The children of 39, 85 and 85 are both bigger than 39. Same for 16 whose children are 92 and 76, which are both larger than 16.

2. Drawn by hand.

3. The final valid max heap would correspond to the array: [92,85,85,76,49,30,49,16,39,15,21,7,29,31,28]



## Question 2. 
Consider the following questions on the $MAX-HEAPIFY$ operation.
### Question 2a.

In the pseudocode of $MAX-HEAPIFY$ (Cormen et al., p.154, or you can view it [here](https://drive.google.com/open?id=1e_3jsX4-qQCfZXKMok_T6LPFh9FwtmT5)), what does A.heap-size mean and what is the idea behind the local variable largest? 


<!-- grade_id: cell-06106b81a909b9e6, points: 0, grade: true, solution: true -->
A.heap-size refers to the size of the (sub)tree that we are currently looking at. The idea behind the local variable largest is that we sort the subtree (treat it like a tree) according to the largest variable in this subtree - the largest local variable.

### Question 2b.
The functions $LEFT(i)$ and $RIGHT(i)$, lines 1 and 2 in the $MAX-HEAPIFY$ pseudocode, return the array index of the left and right child, respectively, of a node in a binary tree. From reading Section 6.1, you know that the input to both functions is an integer number, $i$, which corresponds to the array index of the parent node in the array. Review Section 6.1 for more information. Write a Python implementation of the functions $LEFT(i)$ and $RIGHT(i)$ by filling in the cells below.

<!-- grade_id: cell-2efd4017321c8d48, points: 0, grade: false, solution: true -->
```python
def left(i):
    leftie = (2*i)
    return leftie
    raise NotImplementedError()
```

<!-- grade_id: cell-507ddbab62e6a32c, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

<!-- grade_id: cell-b7e5ceeacaeaf1b8, points: 0, grade: false, solution: true -->
```python
def right(i):
    rightie = (2*i)+1
    return rightie
    raise NotImplementedError()
```

<!-- grade_id: cell-c24867515653788a, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Question 2c.
Write a Python implementation of the MAX-HEAPIFY operation using the pseudocode above, and your newly written functions, `left` and `right`.

<!-- grade_id: cell-52289253a243341b, points: 0, grade: false, solution: true -->
```python
def heapify(heap, i):
    """
    Inputs:
    - heap: a list of floats. Assume that the heap size is the length of the heap.
    
    No output is needed. This function should modify (if necessary) heap in-place.
    """
    
    largest = i # initialize largest variable
    heapsize = len(heap)-1
    
    l = left(i)
    # print(l)
    # print(heap[l])
    r = right(i)
    # print(r)
    # print(heap[r])
    
    if l <= heapsize:
        if heap[l] > heap[i]:
            largest = l
    if r <= heapsize:
        if heap[r] > heap[largest]:
            largest = r
    if largest != i:
        temp = heap[i]
        heap[i] = heap[largest]
        heap[largest] = temp
        heapify(heap, largest)
    
    # raise NotImplementedError()
```

```python
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
heapify(A,0)
A
```

Output:
```
[85, 85, 49, 16, 76, 7, 49, 92, 39, 15, 21, 30, 29, 31, 28]
```

Note: I've tried de-bugging but can't find my error. The function seems to work fine and follows the pseudocode (to the best of my ability) but still doesn't return quite what we want. Any insights much appreciated!

<!-- grade_id: cell-7f63745ff34b881b, points: 1, grade: true, solution: false -->
```python
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
heapify(A,0)
assert(A == [85, 49, 85, 16, 39, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28])

```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-71-f270faf6e16e> in <module>
      1 A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
      2 heapify(A,0)
----> 3 assert(A == [85, 49, 85, 16, 39, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28])

AssertionError:
```

## Question 3. 
Next, write a Python implementation of the BUILD_MAX_HEAP operation using the pseudocode provided in Section 6.3 of Cormen et. al. Test your Python implementation using the array in problem 1, and make sure your Python codes produce a valid max heap.

<!-- grade_id: cell-9237da6268a5e3eb, points: 0, grade: false, solution: true -->
```python
def build_max_heap(A):
    """
    Input:
    - A: a list of floats.
    
    No output is needed. The function should turn A into a valid max heap, in-place.
    """
    
    heap_size = len(A)-1
    
    for i in range(heap_size//2, 0, -1):
        heapify(A, i)
        
    # raise NotImplementedError()
```

```python
A = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]
build_max_heap(A)
A
```

Output:
```
[4, 16, 9, 14, 8, 1, 10, 2, 3, 7]
```

<!-- grade_id: cell-5dc147b837da3f7e, points: 1, grade: true, solution: false -->
```python
A = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]
build_max_heap(A)
assert(A == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1])
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-74-ee91f00716e1> in <module>
      1 A = [4, 1, 3, 2, 16, 9, 10, 14, 8, 7]
      2 build_max_heap(A)
----> 3 assert(A == [16, 14, 10, 8, 7, 9, 3, 2, 4, 1])

AssertionError:
```

## Question 4. 

Lastly, write Python implementations of the $MIN-HEAPIFY$ and $BUILD-MIN-HEAP$ operations for a min heap data structure. You can use your $MAX-HEAPIFY$ and $BUILD-MAX-HEAP$ Python function as models, just remember that the latter two functions support operations for the max heap data structure. Test your Python implementation of $BUILD-MIN-HEAP$ using the array in problem 1, and make sure your Python codes produce a valid min heap. 

<!-- grade_id: cell-d86943f127876fc8, points: 0, grade: false, solution: true -->
```python
def min_heapify(heap, i):
    """
    Inputs:
    - heap: a list of floats. Assume that the heap size is the length of the heap.
    
    No output is needed. This function should modify (if necessary) heap in-place.
    """
    
    smallest = i # initialize smallest variable
    heapsize = len(heap)-1
    
    l = left(i)
    # print(l)
    # print(heap[l])
    r = right(i)
    # print(r)
    # print(heap[r])
    
    if l <= heapsize:
        if heap[l] < heap[i]:
            smallest = l
    if r <= heapsize:
        if heap[r] < heap[smallest]:
            smallest = r
    if smallest != i:
        temp = heap[i]
        heap[i] = heap[smallest]
        heap[smallest] = temp
        min_heapify(heap, smallest)

    # raise NotImplementedError()
```

<!-- grade_id: cell-6efa45e3956ad0b8, points: 0, grade: false, solution: true -->
```python
def build_min_heap(A):
    """
    Input:
    - A: a list of floats.
    
    No output is needed. The function should turn A into a valid min heap, in-place.
    """
    
    heap_size = len(A)-1
    
    for i in range(heap_size//2, 0, -1):
        heapify(A, i)
        
    
    # raise NotImplementedError()
```

<!-- grade_id: cell-7e64941402f819c2, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

```python
# Testing
A = [39, 85, 85, 16, 49, 7, 49, 92, 76, 15, 21, 30, 29, 31, 28]
build_min_heap(A)
A
```

Output:
```
[39, 92, 85, 85, 76, 30, 49, 28, 49, 15, 21, 7, 29, 31, 16]
```

Again, guidance would be much appreciated.