```python
NAME = "Magali"
COLLABORATORS = "NA"
```

---

# CS110 Pre-class Work 7.1

## Part A. Median Heap (watch this [video explanation](https://www.youtube.com/watch?v=756_8C9YBZQ&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=5&t=0s) or read this [description](https://stackoverflow.com/a/15319593/7946759))

Throughout this pre-class work, please use the following definition of median: the median of a list of numbers is the one in the middle of the list when the list is ordered. When such the middle element can’t be determined (i.e., in a list of even length), the average of the two middle elements is the median. For example, 5 is the median of [-1,2,4,5,8,10,12], and (5+7)/2=6 is the median of [1,2,3,5,7,8,10,11].

Using the idea from Lesson 3.2, we can use a pair of heaps to create a data structure which allows fast access to the median. Use the heapq module in python to create both a max-heap and a min-heap. Note that by default, the heapq module in python only creates min-heaps, but if we multiply elements by -1 when we store them, then we can also create max-heaps.


## Question 1.
Write a function `add_to_median_heap(minh, maxh, elem)`. It must accept a min heap, a max heap, and an element to add.


<!-- grade_id: cell-50cb08de70712ee7, points: 0, grade: false, solution: true -->
```python
import heapq

def add_to_median_heap(minh, maxh, elem):
    # find the current median (call median function)
    cmed = median(minh, maxh)
    # if both queues are empty
    if cmed == float('inf'):
        cmed = elem
    # compare the element and median
    if elem <= cmed: # is the element <= the current median?
        # if len(maxh) > len(minh):
            # heapq.heappop
        heapq.heappush(maxh, elem) # insert into max heap
    if elem >= cmed: # is the element >= the current median?
        heapq.heappush(minh, elem) # insert into the min heap
```

## Question 2
Write a function `median(minh, maxh)`. It must return the median element.


<!-- grade_id: cell-43ce866dd1dc2d5b, points: 0, grade: false, solution: true -->
```python
def median(minh, maxh):
    # get lengths of heaps
    lminh = len(minh)
    lmaxh = len(maxh)
    # when the min-heap contains one more element than the max-heap, 
    # the median is in the top of the min-heap.
    if lminh > lmaxh:
        med = heapq.heappop(minh)
        return med
    # when the max-heap contains one more element than the min-heap, 
    # the median is in the top of the max-heap.
    elif lmaxh > lminh:
        med = heapq.heappop(maxh)
        return med
    # when both heaps contain the same number of elements, 
    # the total number of elements is even and the median is 
    # the mean of the two middle elements
    else: # lmaxh == lminh:
        if lmaxh == 0:
            return float('inf')
        else:
            med1 = heapq.heappop(minh) 
            med2 = heapq.heappop(maxh) 
            med = (med1 + med2)/2
            return med
```

<!-- grade_id: cell-673faa69bf330eda, points: 0, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

## Question 3.

Uncomment and run the testing code given below to test your functions. It should print out numbers ranging from 1 to 50, in that order.

```python
minh = []
# heapq.heapify(minh)
maxh = []
# heapq.heapify(maxh)
for a in range(1,100,2):
    add_to_median_heap(minh, maxh, a)
    print(median(minh, maxh))  
```

Output:
```
1.0
3.0
5.0
7.0
9.0
11.0
13.0
15.0
17.0
19.0
21.0
23.0
25.0
27.0
29.0
31.0
33.0
35.0
37.0
39.0
41.0
43.0
45.0
47.0
49.0
51.0
53.0
55.0
57.0
59.0
61.0
63.0
65.0
67.0
69.0
71.0
73.0
75.0
77.0
79.0
81.0
83.0
85.0
87.0
89.0
91.0
93.0
95.0
97.0
99.0
```

The instructions say that it should print numbers until 50 but here it prints until 100, according to the defined range function (from 0 to 100, skip by 2 / odd numbers). 

## Question 4.
What’s the worst case complexity to build a median heap using `add_to_median_heap`?

<!-- grade_id: cell-c8e33ebdc8aaf328, points: 0, grade: true, solution: true -->
The worst-case complexity to build a median heap is equivalent to the complexity it takes to push an element in: O(n log n), where n is roughly half of all the elements. I believe that this is the same as the average case complexity.

## Question 5.
What’s the worst case complexity of `median`?

<!-- grade_id: cell-c061920faa337890, points: 0, grade: true, solution: true -->
The worst-case complexity of median is equivalent to the complexity it takes to pop an element: O(n log n) (and some constant time when calculating the mean of the two), where n is roughly half of all the elements. I believe that this is the same as the average case complexity.

## Question 6.

How does this way of finding the median compare with the vanilla way of sorting the list and pick the middle element? Use arguments based on efficiency or clearity of the respective algorithms.

<!-- grade_id: cell-1f31afecda9730a7, points: 0, grade: true, solution: true -->
The most efficient sorting algorithms have an average case complexity of O(n log n) too (merge sort, heap sort, quick sort).  Nevertheless, n is defined here as all of the elements, unlike with the heaps. Picking the middle element is constant time (it requires just the division of the length by two).  Hence, while the "vanilla" way of sorting is O(n log n), the heaps median algorithm complexity can be represented by O(n/2 log n/2) + O(n/2 log n/2). 

## [Optional] Question 7.

Is it possible to extend this idea to any percentile? If it is, then write code to do so. If it’s not possible, prove why it is not possible.

<!-- grade_id: cell-4d0c197e6d484338, points: 0, grade: true, solution: true -->
What does it mean: any percentile?  As in the median is at the 50th percentile and we are wondering whether we can do a min-/max-heap for a number that's at, say, the 75th percentile?

## Part B. Quickselect

Quicksort can be modified to find the $k$-th smallest element in an unordered list. This is known as quickselect. It does this by choosing a partition (as in quicksort). Once the list has been partitioned then we know how many elements lie to the left and to the right of the partition. This allows us to recursively call quickselect on the correct sublist.

## Question 1.

Write a function `qselect(lst, k)`, which takes both a list and an index $k$. The function must then return the $k$-th smallest item in the list.


I'm not sure I understand the prompt correctly.  Is this what is meant by quickselect?  If not, how can we know that the element is the k_th smallest if the list is unordered?  Do we first need to order the list?

<!-- grade_id: cell-fb9c2bc9d5b4bc92, points: 0, grade: false, solution: true -->
```python
def qselect(lst,k):
    F = quick_sort(lst, 0, len(lst)-1)
    elem = F[k]
    return elem
```

```python
def quick_sort(A,p,r):
    if p < r:
        q = partition(A,p,r)
        quick_sort(A,p,q-1)
        quick_sort(A,q+1,r)
    return A
```

```python
def partition(A,p,r):
    """
    Assume r<len(A) and p>=0
    """
    x = A[r]
    i = p-1
    for j in range(p, r):
        if A[j] <= x:
            i = i + 1
            temp = A[i]
            A[i] = A[j]
            A[j] = temp
    tempr = A[r]
    A[r] = A[i+1]
    A[i+1] = tempr
    return i+1
```

<!-- grade_id: cell-85a25ed6a1f62a66, points: 0, grade: true, solution: false -->
```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst1 = list(range(100))
random.shuffle(lst1)
lst2 = []
for a in range(100):
    lst2.append(qselect(lst1, a))
assert(lst2 == list(range(100)))
```

## Question 2.
Uncomment and run the testing code given below to test your function. It should print out integers from 0 to 99.


```python
import random
random.seed(123) # introducing a seed for reproducibility purposes
lst = list(range(100))
random.shuffle(lst)
for a in range(100):     
    print(qselect(lst, a))

```

Output:
```
0
1
2
3
4
5
6
7
8
9
10
11
12
13
14
15
16
17
18
19
20
21
22
23
24
25
26
27
28
29
30
31
32
33
34
35
36
37
38
39
40
41
42
43
44
45
46
47
48
49
50
51
52
53
54
55
56
57
58
59
60
61
62
63
64
65
66
67
68
69
70
71
72
73
74
75
76
77
78
79
80
81
82
83
84
85
86
87
88
89
90
91
92
93
94
95
96
97
98
99
```

## Question 3.

Write down the recurrence relation for your code.

<!-- grade_id: cell-24144607b3ffade3, points: 0, grade: true, solution: true -->
T(N) = N + T(N-1) 

or 

T(N) = N + 2T(N/2)

## Question 4.

Solve the recurrence relation for quickselect in the best case.


<!-- grade_id: cell-c125f4742933e820, points: 0, grade: true, solution: true -->
I'll take the second one to be the best case. 

T(N) = N + 2T(N/2)

No idea how to go about it.  But I think this should be equal to O(N log (N)).  I'll be going to OH to ask about this once I receive a review.

## Question 5.

Solve the recurrence relation for quickselect in the worst case.

<!-- grade_id: cell-852c3a71d8c2d7c1, points: 0, grade: true, solution: true -->
T(N) = N + T(N-1) 
 

T(N-1) = (N-1) + T(N-2) 


T(N-2) = (N-2) + T(N-3)


...

 

T(2) = (2) + T(1) = 2

 

T(1) = 1 + T(0) = 0 ? 

 

T(0) = (0) + T(N-1) => T(1) = 0



=> T(N) = N + (N-1) + (N-2) ... + 2 + 0


= (N^2)/2

= O(N^2)